<img src="Archivos/miad4.png" width=800x>

# Explorar y modificar bases de datos utilizando `pandas`

En este taller filtrarás y modificarás bases de datos por medio de sus índices al resolver diversos ejercicios. 

## Habilidades en práctica

Al realizar este taller podrás revisar tu progreso para:

**1.** Indexar un `DataFrame` empleando índices unitarios o múltiples. <br>
**2.** Aplicar los distintos tipos de filtrado utilizando `pandas`.<br>
**3.** Emplear algunos métodos de `pandas` para modificar una base de datos. <br>

## Instrucciones

En cada uno de los siguientes ejercicios deberás escribir el código solicitado estrictamente en las celdas indicadas para ello, teniendo en cuenta las siguientes recomendaciones:

* No crear, eliminar o modificar celdas de este Notebook (salvo lo que se te indique), pues puede verse afectado el proceso de calificación automática.

* La calificación se realiza de manera automática con datos diferentes a los proporcionados en este taller. Por consiguiente, tu código debe funcionar para diferentes instancias de cada uno de los ejercicios; una instancia hace referencia a los posibles valores de los parámetros.

* La calificación de cada ejercicio depende del valor que retorne la función especificada en su enunciado. Por lo tanto, aunque implementes funciones adicionales, es escencial que utilices los nombres propuestos en los enunciados de los ejercicios para implementar la función definitiva.

## Ejercicios
En la siguente celda encuentras declarados los paquetes necesarios para el desarollo de este taller.

```python
# Esta celda no es modificable

import numpy as np
import pandas as pd
```

### Ejercicio 1

Junto a este Notebook encontrarás un archivo llamado `"pooled_data_public.dta"` que contiene la base de datos completa del BID. Este archivo, al ser de tipo *.dta*, requiere de una función de `pandas` para ser importado a Python.

    read_stata(file = "Archivos/pooled_data_public.dta")

Implementa una función llamada `importar_e_indexar` que no reciba parámetros, importe los datos del archivo `"pooled_data_public.dta"` y retorne un `DataFrame` cuyo índice (`index`) sea la columna `'id'` de la base de datos sin incluirla en las columnas del `DataFrame` resultante.

La función debe retornar un `DataFrame`.

```python
df1 = pd.read_stata("pooled_data_public.dta")
```

Output:
```
/usr/local/lib/python3.7/dist-packages/pandas/io/stata.py in read_stata(filepath_or_buffer, convert_dates, convert_categoricals, index_col, convert_missing, preserve_dtypes, columns, order_categoricals, chunksize, iterator)
   1928 
   1929     try:
-> 1930         data = reader.read()
   1931     finally:
   1932         reader.close()

/usr/local/lib/python3.7/dist-packages/pandas/io/stata.py in read(self, nrows, convert_dates, convert_categoricals, index_col, convert_missing, preserve_dtypes, columns, order_categoricals)
   1602         read_lines = min(nrows, self.nobs - self._lines_read)
   1603         data = np.frombuffer(
-> 1604             self.path_or_buf.read(read_len), dtype=dtype, count=read_lines
   1605         )
   1606 

ValueError: buffer is smaller than requested size
```

```python
df1.head(2)
```

Output:
```
id country_code  ... nps_medio_website  nps_medio_otro
0  1000060.0      Bahamas  ...                No              No
1  1000734.0      Bahamas  ...                No              No

[2 rows x 129 columns]
```

<!-- grade_id: ej1-sol, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def importar_e_indexar():
    df1 = pd.read_stata("pooled_data_public.dta")
    df1=df1.set_index('id')
    return df1
```

```python
df1 = pd.read_stata("pooled_data_public.dta")
df1['id'] = df1['id'].astype(int)
df1=df1.set_index('id')
```

<!-- grade_id: ej1-test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Base variables
import inspect

nombre_columnas = ['country_code', 'country_name', 'weight', 'weight_pop', 'pop', 'recordeddate', 'genero', 'edad', 'educ', 'estcivil', 'nrhogar', 'nrhogar_antes', 'hayninhos', 'ninos_5', 'nrninhos', 'matriculados', 'esc_publica', 'grado', 'haymayores', 'diasmercado', 'lugarmercado', 'comprotodo', 'nocompra_pan', 'nocompra_vegetal', 'nocompra_carne', 'nocompra_envasados', 'nocompra_bebidas', 'nocompra_limpieza', 'nocompra_medicinas', 'nocompra_otro', 'nrnegocios', 'esperarfila', 'subaprecios', 'hambre', 'menos_saludable', 'vivienda', 'probvivienda', 'probvivienda_prox', 'ayuda_alquiler', 'sufialimento', 'sufirecursos', 'gasto_emergencia', 'monto_gasto', 'pre_cash', 'new_cash', 'coima', 'tiene_banco', 'tiene_tarjdebito', 'tiene_segsaludpub', 'tiene_segsaludpriv', 'recibeprestamos', 'darprestamos', 'remesas_enero', 'remesas_origen', 'remesas_persona', 'recibio_remesas', 'lm_perdioempleo', 'lm_cerronegocio_gob', 'lm_cerronegocio_dem', 'lm_asistmedica', 'lm_perdioingrarriendo', 'lm_dejoasistencia', 'lm_freq', 'recibio_ayudaemp', 'compromiso_futurotrabajo', 'empleo_formal', 'actividad_empleado', 'actividad_negocio', 'actividad_estudio', 'trabajocasa', 'diastrabajo', 'horastrabajo', 'sector', 'redujo_ingreso', 'ingreso_jan2020', 'ingreso_apr2020', 'retorno_ingreso', 'sintoma_ninguno', 'sintoma_garganta', 'sintoma_gripe', 'sintoma_sarpullido', 'sintoma_fiebre', 'sintoma_tos', 'conoce_sintomas', 'infectados_sinsintomas', 'sinsintomas_contagiar', 'visito_trabajo', 'visito_hospital', 'visito_mercado', 'visito_banco', 'visito_amigos', 'visito_familiares', 'visito_ejercitar', 'visito_iglesia', 'comotransmite', 'sabe_comotransmite', 'distsoc', 'recom_trabajo', 'recom_hospital', 'recom_mercado', 'recom_banco', 'recom_amigos', 'recom_familiares', 'recom_ejercicio', 'recom_iglesia', 'haceds_hogar', 'haceds_vecinos', 'haceds_jovenes', 'haceds_mayores', 'haceds_otropais', 'haceds_emplmercados', 'medios_noti_redessociales', 'medios_noti_chat', 'medios_noti_periodicos', 'medios_noti_tv', 'medios_noti_radio', 'medios_covid_redessociales', 'medios_covid_chat', 'medios_covid_periodicos', 'medios_covid_tv', 'medios_covid_radio', 'nps_medio_redessoc', 'nps_medio_sms', 'nps_medio_periodico', 'nps_medio_tv', 'nps_medio_radio', 'nps_medio_website', 'nps_medio_otro']
tipo_columnas = ['CategoricalDtype', 'dtype[object_]', 'dtype[float32]', 'dtype[float32]', 'dtype[float64]', 'dtype[object_]', 'CategoricalDtype', 'dtype[int16]', 'CategoricalDtype', 'CategoricalDtype', 'dtype[int8]', 'dtype[int8]', 'CategoricalDtype', 'CategoricalDtype', 'dtype[float64]', 'dtype[float64]', 'CategoricalDtype', 'dtype[float64]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'dtype[float32]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'dtype[float64]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'dtype[float64]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'dtype[int8]', 'dtype[int8]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'dtype[float32]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'dtype[object_]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype']

# Caso 1: no existe la función.
try:
    source = inspect.getsource(importar_e_indexar)
    assert type(importar_e_indexar) == type(lambda:None)
except:
    raise NotImplementedError("No existe una función llamada importar_e_indexar.",)
    
# Caso 2: la función es interrumpida por errores durante su ejecución.
try:
    resultado = importar_e_indexar()
except:
    raise RuntimeError("Tu función produce un error al ejecutarse.")

# Caso 3: no retorna un DataFrame
assert type(resultado) == pd.DataFrame, f"Tu función debe retornar un valor de tipo {pd.DataFrame.__name__}."

# Caso 4: no utiliza el nombre del archivo
assert source.find("pooled_data_public.dta") >= 0, "Tu código no hace referencia al archivo 'pooled_data_public.dta'."

# Caso 5: retorna el dataframe con la columna 'id'
assert "id" not in resultado.columns, "Tu función retorna un DataFrame que incluye la columna 'id'. Debes convertir esta columna en el index del DataFrame."

# Caso 6: retorna un dataframe con cantidad de columnas errada
assert resultado.shape[1] == 128, "Tu función retorna un DataFrame con cantidad de columnas errada."

# Caso 7: devuelve un dataframe cuyas columnas tienen nombre distinto de lo esperado
assert [*resultado.columns] == nombre_columnas, "Tu función retorna un DataFrame con columnas cuyos nombres no coinciden con los de los datos del archivo."

# Caso 8: devuelve un dataframe cuyas columnas tienen tipo distinto de lo esperado
assert [i.__class__.__name__ for i in resultado.dtypes] == tipo_columnas, "Tu función retorna un DataFrame con columnas cuyos tipos no coinciden con los de los datos del archivo."

# Caso 9: devuelve un dataframe con cantidad de filas errada
assert resultado.shape[0] == 230509, "Tu función retorna un DataFrame con cantidad de filas errada."

# Caso 10: devuelve un dataframe con indice del mismo tipo que por defecto ('int64')
assert resultado.index.dtype.name != 'int64', "Tu función retorna un DataFrame con índice del mismo tipo que el que se crea por defecto en pandas. Asegurate de estar asignando la columna 'id' al índice."

# Caso 11: devuelve un dataframe con índice de tipo errado
assert resultado.index.dtype.name == 'float64', "Tu función retorna un DataFrame con índice de tipo distinto del esperado. Asegurate de estar usando la columna 'id' sin editarla."

# Caso 12: el índice no tiene el nombre de la columna 'id'
assert resultado.index.name == 'id', "Tu función retorna un DataFrame con índice de nombre distinto de la columna 'id'."

# Caso 13: primer o última posicion del índice distintas de lo esperado
assert resultado.index[0] == 1000060.0, "Tu función retorna un DataFrame que es errado por lo menos en la primera posición de su índice"
assert resultado.index[len(resultado.index) - 1] == 17003990.0, "Tu función retorna un DataFrame que es errado por lo menos en la última posición de su índice"

# Mensaje de felicitaciones
print("Felicidades, realizaste este ejercicio correctamente.")
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-6-d143d168ea2b> in <module>()
     36 
     37 # Caso 8: devuelve un dataframe cuyas columnas tienen tipo distinto de lo esperado
---> 38 assert [i.__class__.__name__ for i in resultado.dtypes] == tipo_columnas, "Tu función retorna un DataFrame con columnas cuyos tipos no coinciden con los de los datos del archivo."
     39 
     40 # Caso 9: devuelve un dataframe con cantidad de filas errada

AssertionError: Tu función retorna un DataFrame con columnas cuyos tipos no coinciden con los de los datos del archivo.
```

### Ejercicio 2

Implementa una función llamada `indice_sencillo_a_multiple` que no reciba parámetros, tome el `DataFrame` resultante del [Ejercicio 1](#Ejercicio-1) y modifique el `DataFrame` de manera que su índice (`index`) consista de tuplas compuestas por el valor del índice original y el valor correspondiente en la columna `"country_code"`. No incluya la columna `"country_code"` en el `DataFrame` resultante.

La función debe retornar un `DataFrame`.

<!-- grade_id: ej2-sol, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def indice_sencillo_a_multiple():
    df2 = importar_e_indexar()
    df2.set_index(keys = "country_code", inplace = True,
                   append = True, drop = True)
    return df2
```

<!-- grade_id: ej2-test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Base variables
import inspect

nombre_columnas = ['country_name', 'weight', 'weight_pop', 'pop', 'recordeddate', 'genero', 'edad', 'educ', 'estcivil', 'nrhogar', 'nrhogar_antes', 'hayninhos', 'ninos_5', 'nrninhos', 'matriculados', 'esc_publica', 'grado', 'haymayores', 'diasmercado', 'lugarmercado', 'comprotodo', 'nocompra_pan', 'nocompra_vegetal', 'nocompra_carne', 'nocompra_envasados', 'nocompra_bebidas', 'nocompra_limpieza', 'nocompra_medicinas', 'nocompra_otro', 'nrnegocios', 'esperarfila', 'subaprecios', 'hambre', 'menos_saludable', 'vivienda', 'probvivienda', 'probvivienda_prox', 'ayuda_alquiler', 'sufialimento', 'sufirecursos', 'gasto_emergencia', 'monto_gasto', 'pre_cash', 'new_cash', 'coima', 'tiene_banco', 'tiene_tarjdebito', 'tiene_segsaludpub', 'tiene_segsaludpriv', 'recibeprestamos', 'darprestamos', 'remesas_enero', 'remesas_origen', 'remesas_persona', 'recibio_remesas', 'lm_perdioempleo', 'lm_cerronegocio_gob', 'lm_cerronegocio_dem', 'lm_asistmedica', 'lm_perdioingrarriendo', 'lm_dejoasistencia', 'lm_freq', 'recibio_ayudaemp', 'compromiso_futurotrabajo', 'empleo_formal', 'actividad_empleado', 'actividad_negocio', 'actividad_estudio', 'trabajocasa', 'diastrabajo', 'horastrabajo', 'sector', 'redujo_ingreso', 'ingreso_jan2020', 'ingreso_apr2020', 'retorno_ingreso', 'sintoma_ninguno', 'sintoma_garganta', 'sintoma_gripe', 'sintoma_sarpullido', 'sintoma_fiebre', 'sintoma_tos', 'conoce_sintomas', 'infectados_sinsintomas', 'sinsintomas_contagiar', 'visito_trabajo', 'visito_hospital', 'visito_mercado', 'visito_banco', 'visito_amigos', 'visito_familiares', 'visito_ejercitar', 'visito_iglesia', 'comotransmite', 'sabe_comotransmite', 'distsoc', 'recom_trabajo', 'recom_hospital', 'recom_mercado', 'recom_banco', 'recom_amigos', 'recom_familiares', 'recom_ejercicio', 'recom_iglesia', 'haceds_hogar', 'haceds_vecinos', 'haceds_jovenes', 'haceds_mayores', 'haceds_otropais', 'haceds_emplmercados', 'medios_noti_redessociales', 'medios_noti_chat', 'medios_noti_periodicos', 'medios_noti_tv', 'medios_noti_radio', 'medios_covid_redessociales', 'medios_covid_chat', 'medios_covid_periodicos', 'medios_covid_tv', 'medios_covid_radio', 'nps_medio_redessoc', 'nps_medio_sms', 'nps_medio_periodico', 'nps_medio_tv', 'nps_medio_radio', 'nps_medio_website', 'nps_medio_otro']
tipo_columnas = ['dtype[object_]', 'dtype[float32]', 'dtype[float32]', 'dtype[float64]', 'dtype[object_]', 'CategoricalDtype', 'dtype[int16]', 'CategoricalDtype', 'CategoricalDtype', 'dtype[int8]', 'dtype[int8]', 'CategoricalDtype', 'CategoricalDtype', 'dtype[float64]', 'dtype[float64]', 'CategoricalDtype', 'dtype[float64]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'dtype[float32]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'dtype[float64]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'dtype[float64]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'dtype[int8]', 'dtype[int8]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'dtype[float32]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'dtype[object_]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype']

# Caso 1: no existe la función.
try:
    source = inspect.getsource(indice_sencillo_a_multiple)
    assert type(indice_sencillo_a_multiple) == type(lambda:None)
except:
    raise NotImplementedError("No existe una función llamada indice_sencillo_a_multiple.",)
    
# Caso 2: la función es interrumpida por errores durante su ejecución.
try:
    resultado = indice_sencillo_a_multiple()
except:
    raise RuntimeError("Tu función produce un error al ejecutarse.")

# Caso 3: no retorna un DataFrame
assert type(resultado) == pd.DataFrame, f"Tu función debe retornar un valor de tipo {pd.DataFrame.__name__}."

# Caso 4: no ejecuta la función importar_e_indexar
assert source.find("importar_e_indexar()") >= 0, "Tu código no hace referencia a la función importar_e_indexar. Asegurate de ejecutarla en tu función."

# Caso 5: retorna el dataframe con la columna 'id'
assert "id" not in resultado.columns, "Tu función retorna un DataFrame que incluye la columna 'id'. Esta columna debes haberla convertido en el index sencillo del DataFrame."

# Caso 6: retorna el dataframe con la columna 'country_code'
assert "country_code" not in resultado.columns, "Tu función retorna un DataFrame que incluye la columna 'country_code'. Esta columna debes haberla unido al index del DataFrame para formar un índice múltiple."

# Caso 7: retorna un dataframe con cantidad de columnas errada
assert resultado.shape[1] == 127, "Tu función retorna un DataFrame con cantidad de columnas errada."

# Caso 8: devuelve un dataframe cuyas columnas tienen nombre distinto de lo esperado
assert [*resultado.columns] == nombre_columnas, "Tu función retorna un DataFrame con columnas cuyos nombres no coinciden con los de los datos del archivo."

# Caso 9: devuelve un dataframe cuyas columnas tienen tipo distinto de lo esperado
assert [i.__class__.__name__ for i in resultado.dtypes] == tipo_columnas, "Tu función retorna un DataFrame con columnas cuyos tipos no coinciden con los de los datos del archivo."

# Caso 10: devuelve un dataframe con cantidad de filas errada
assert resultado.shape[0] == 230509, "Tu función retorna un DataFrame con cantidad de filas errada."

# Caso 11: devuelve un dataframe con indice no multiple
assert type(resultado.index) == pd.MultiIndex, "Tu función retorna un DataFrame con índice no múltiple."

# Caso 12: devuelve un dataframe con primer nivel de indice multiple del mismo tipo que por defecto ('int64')
assert resultado.index.get_level_values(0).dtype.name != 'int64', "Tu función retorna un DataFrame con índice múltiple cuyo primer nivel es del mismo tipo que el que se crea por defecto en pandas. Asegurate de estar asignando la columna 'id' al índice."

# Caso 13: devuelve un dataframe con primer nivel de indice multiple de tipo errado
assert resultado.index.get_level_values(0).dtype.name == 'float64', "Tu función retorna un DataFrame con índice múltiple cuyo primer nivel es de tipo distinto del esperado. Asegurate de estar usando la columna 'id' sin editarla."

# Caso 14: devuelve un dataframe con segundo nivel de indice multiple de tipo errado
assert resultado.index.get_level_values(1).dtype.name == 'category', "Tu función retorna un DataFrame con índice múltiple cuyo segundo nivel es de tipo distinto del esperado. Asegurate de estar usando la columna 'country_code' sin editarla."

# Caso 15: el multiíndice no tiene el nombre de la columna 'id' y 'country_code'
assert [*resultado.index.names] == ['id', 'country_code'], "Tu función retorna un DataFrame con índice de nombres distintos de las columna 'id' y 'country_code'."

# Caso 16: primera o última posicion del índice distintas de lo esperado
assert resultado.index[0] == (1000060.0, 'Bahamas'), "Tu función retorna un DataFrame que es errado por lo menos en la primera posición de su índice"
assert resultado.index[len(resultado.index) - 1] == (17003990.0, 'Uruguay'), "Tu función retorna un DataFrame que es errado por lo menos en la última posición de su índice"


# Mensaje de felicitaciones
print("Felicidades, realizaste este ejercicio correctamente.")
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-8-b6353ed752a1> in <module>()
     39 
     40 # Caso 9: devuelve un dataframe cuyas columnas tienen tipo distinto de lo esperado
---> 41 assert [i.__class__.__name__ for i in resultado.dtypes] == tipo_columnas, "Tu función retorna un DataFrame con columnas cuyos tipos no coinciden con los de los datos del archivo."
     42 
     43 # Caso 10: devuelve un dataframe con cantidad de filas errada

AssertionError: Tu función retorna un DataFrame con columnas cuyos tipos no coinciden con los de los datos del archivo.
```

### Ejercicio 3

Implementa una función llamada `sin_columnas` que no reciba parámetros, tome el `DataFrame` resultante del [Ejercicio 1](#Ejercicio-1) y lo retorne sin las columnas `"country_name"`, `"grado"` y `"coima"`.

Debes hacer uso del método `drop`.

La función debe retornar un `DataFrame`.

<!-- grade_id: ej3-sol, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def sin_columnas():
    df3 = importar_e_indexar()
    return df3.drop(columns = ["country_name", "grado", "coima"], inplace = False)
```

<!-- grade_id: ej3-test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Base variables
import inspect

nombre_columnas = ['country_code', 'weight', 'weight_pop', 'pop', 'recordeddate', 'genero', 'edad', 'educ', 'estcivil', 'nrhogar', 'nrhogar_antes', 'hayninhos', 'ninos_5', 'nrninhos', 'matriculados', 'esc_publica', 'haymayores', 'diasmercado', 'lugarmercado', 'comprotodo', 'nocompra_pan', 'nocompra_vegetal', 'nocompra_carne', 'nocompra_envasados', 'nocompra_bebidas', 'nocompra_limpieza', 'nocompra_medicinas', 'nocompra_otro', 'nrnegocios', 'esperarfila', 'subaprecios', 'hambre', 'menos_saludable', 'vivienda', 'probvivienda', 'probvivienda_prox', 'ayuda_alquiler', 'sufialimento', 'sufirecursos', 'gasto_emergencia', 'monto_gasto', 'pre_cash', 'new_cash', 'tiene_banco', 'tiene_tarjdebito', 'tiene_segsaludpub', 'tiene_segsaludpriv', 'recibeprestamos', 'darprestamos', 'remesas_enero', 'remesas_origen', 'remesas_persona', 'recibio_remesas', 'lm_perdioempleo', 'lm_cerronegocio_gob', 'lm_cerronegocio_dem', 'lm_asistmedica', 'lm_perdioingrarriendo', 'lm_dejoasistencia', 'lm_freq', 'recibio_ayudaemp', 'compromiso_futurotrabajo', 'empleo_formal', 'actividad_empleado', 'actividad_negocio', 'actividad_estudio', 'trabajocasa', 'diastrabajo', 'horastrabajo', 'sector', 'redujo_ingreso', 'ingreso_jan2020', 'ingreso_apr2020', 'retorno_ingreso', 'sintoma_ninguno', 'sintoma_garganta', 'sintoma_gripe', 'sintoma_sarpullido', 'sintoma_fiebre', 'sintoma_tos', 'conoce_sintomas', 'infectados_sinsintomas', 'sinsintomas_contagiar', 'visito_trabajo', 'visito_hospital', 'visito_mercado', 'visito_banco', 'visito_amigos', 'visito_familiares', 'visito_ejercitar', 'visito_iglesia', 'comotransmite', 'sabe_comotransmite', 'distsoc', 'recom_trabajo', 'recom_hospital', 'recom_mercado', 'recom_banco', 'recom_amigos', 'recom_familiares', 'recom_ejercicio', 'recom_iglesia', 'haceds_hogar', 'haceds_vecinos', 'haceds_jovenes', 'haceds_mayores', 'haceds_otropais', 'haceds_emplmercados', 'medios_noti_redessociales', 'medios_noti_chat', 'medios_noti_periodicos', 'medios_noti_tv', 'medios_noti_radio', 'medios_covid_redessociales', 'medios_covid_chat', 'medios_covid_periodicos', 'medios_covid_tv', 'medios_covid_radio', 'nps_medio_redessoc', 'nps_medio_sms', 'nps_medio_periodico', 'nps_medio_tv', 'nps_medio_radio', 'nps_medio_website', 'nps_medio_otro']
tipo_columnas = ['CategoricalDtype', 'dtype[float32]', 'dtype[float32]', 'dtype[float64]', 'dtype[object_]', 'CategoricalDtype', 'dtype[int16]', 'CategoricalDtype', 'CategoricalDtype', 'dtype[int8]', 'dtype[int8]', 'CategoricalDtype', 'CategoricalDtype', 'dtype[float64]', 'dtype[float64]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'dtype[float32]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'dtype[float64]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'dtype[float64]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'dtype[int8]', 'dtype[int8]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'dtype[float32]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'dtype[object_]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype']

# Caso 1: no existe la función.
try:
    source = inspect.getsource(sin_columnas)
    assert type(sin_columnas) == type(lambda:None)
except:
    raise NotImplementedError("No existe una función llamada sin_columnas.",)
    
# Caso 2: la función es interrumpida por errores durante su ejecución.
try:
    resultado = sin_columnas()
except:
    raise RuntimeError("Tu función produce un error al ejecutarse.")

# Caso 3: no retorna un DataFrame
assert type(resultado) == pd.DataFrame, f"Tu función debe retornar un valor de tipo {pd.DataFrame.__name__}."

# Caso 4: no ejecuta la función importar_e_indexar
assert source.find("importar_e_indexar()") >= 0, "Tu código no hace referencia a la función importar_e_indexar. Asegurate de ejecutarla en tu código."

# Caso 5: no ejecuta el método drop de pandas
assert source.find(".drop(") >= 0, "Tu código no hace referencia al método drop. Asegurate de ejecutarlo en tu código."

# Caso 6: retorna el dataframe con la columna 'id'
assert "id" not in resultado.columns, "Tu función retorna un DataFrame que incluye la columna 'id'. Esta columna debes convertirla en el index del DataFrame."

# Caso 7: retorna el dataframe con la columna 'country_name'
assert "country_name" not in resultado.columns, "Tu función retorna un DataFrame que incluye la columna 'country_name'."

# Caso 8: retorna el dataframe con la columna 'grado'
assert "grado" not in resultado.columns, "Tu función retorna un DataFrame que incluye la columna 'grado'."

# Caso 9: retorna el dataframe con la columna 'coima'
assert "coima" not in resultado.columns, "Tu función retorna un DataFrame que incluye la columna 'coima'."

# Caso 10: retorna un dataframe con cantidad de columnas errada
assert resultado.shape[1] == 125, "Tu función retorna un DataFrame con cantidad de columnas errada."

# Caso 11: devuelve un dataframe cuyas columnas tienen nombre distinto de lo esperado
assert [*resultado.columns] == nombre_columnas, "Tu función retorna un DataFrame con columnas cuyos nombres no coinciden con los de los datos del archivo, sin las columnas country_name, grado y coima."

# Caso 12: devuelve un dataframe cuyas columnas tienen tipo distinto de lo esperado
assert [i.__class__.__name__ for i in resultado.dtypes] == tipo_columnas, "Tu función retorna un DataFrame con columnas cuyos tipos no coinciden con los de los datos del archivo, sin las columnas country_name, grado y coima."

# Caso 13: devuelve un dataframe con cantidad de filas errada
assert resultado.shape[0] == 230509, "Tu función retorna un DataFrame con cantidad de filas errada."

# Caso 14: devuelve un dataframe con indice del mismo tipo que por defecto ('int64')
assert resultado.index.dtype.name != 'int64', "Tu función retorna un DataFrame con índice del mismo tipo que el que se crea por defecto en pandas. Asegurate de estar asignando la columna 'id' como índice."

# Caso 15: devuelve un dataframe con índice de tipo errado
assert resultado.index.dtype.name == 'float64', "Tu función retorna un DataFrame con índice de tipo distinto del esperado. Asegurate de estar usando la columna 'id' sin editarla."

# Caso 16: primer o última posicion del índice distintas de lo esperado
assert resultado.index[0] == 1000060.0, "Tu función retorna un DataFrame que es errado por lo menos en la primera posición de su índice"
assert resultado.index[len(resultado.index) - 1] == 17003990.0, "Tu función retorna un DataFrame que es errado por lo menos en la última posición de su índice"


# Mensaje de felicitaciones
print("Felicidades, realizaste este ejercicio correctamente.")
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-10-7039269acae6> in <module>()
     48 
     49 # Caso 12: devuelve un dataframe cuyas columnas tienen tipo distinto de lo esperado
---> 50 assert [i.__class__.__name__ for i in resultado.dtypes] == tipo_columnas, "Tu función retorna un DataFrame con columnas cuyos tipos no coinciden con los de los datos del archivo, sin las columnas country_name, grado y coima."
     51 
     52 # Caso 13: devuelve un dataframe con cantidad de filas errada

AssertionError: Tu función retorna un DataFrame con columnas cuyos tipos no coinciden con los de los datos del archivo, sin las columnas country_name, grado y coima.
```

### Ejercicio 4

Implementa una función llamada `indexar_columnas_numericamente` que no reciba parámetros, tome el `DataFrame` resultante del [Ejercicio 1](#Ejercicio-1) y lo retorne sin el grupo de **columnas** cuya posición (indexada desde 0) sea mayor o igual a 80 y menor o igual a 100.

La función debe retornar un `DataFrame`.

<!-- grade_id: ej4-sol, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def indexar_columnas_numericamente():
    df4 = importar_e_indexar()
    return df4.drop(columns = df4.columns[80:101], inplace = False)
```

<!-- grade_id: ej4-test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Base variables
import inspect

nombre_columnas = ['country_code', 'country_name', 'weight', 'weight_pop', 'pop', 'recordeddate', 'genero', 'edad', 'educ', 'estcivil', 'nrhogar', 'nrhogar_antes', 'hayninhos', 'ninos_5', 'nrninhos', 'matriculados', 'esc_publica', 'grado', 'haymayores', 'diasmercado', 'lugarmercado', 'comprotodo', 'nocompra_pan', 'nocompra_vegetal', 'nocompra_carne', 'nocompra_envasados', 'nocompra_bebidas', 'nocompra_limpieza', 'nocompra_medicinas', 'nocompra_otro', 'nrnegocios', 'esperarfila', 'subaprecios', 'hambre', 'menos_saludable', 'vivienda', 'probvivienda', 'probvivienda_prox', 'ayuda_alquiler', 'sufialimento', 'sufirecursos', 'gasto_emergencia', 'monto_gasto', 'pre_cash', 'new_cash', 'coima', 'tiene_banco', 'tiene_tarjdebito', 'tiene_segsaludpub', 'tiene_segsaludpriv', 'recibeprestamos', 'darprestamos', 'remesas_enero', 'remesas_origen', 'remesas_persona', 'recibio_remesas', 'lm_perdioempleo', 'lm_cerronegocio_gob', 'lm_cerronegocio_dem', 'lm_asistmedica', 'lm_perdioingrarriendo', 'lm_dejoasistencia', 'lm_freq', 'recibio_ayudaemp', 'compromiso_futurotrabajo', 'empleo_formal', 'actividad_empleado', 'actividad_negocio', 'actividad_estudio', 'trabajocasa', 'diastrabajo', 'horastrabajo', 'sector', 'redujo_ingreso', 'ingreso_jan2020', 'ingreso_apr2020', 'retorno_ingreso', 'sintoma_ninguno', 'sintoma_garganta', 'sintoma_gripe', 'recom_amigos', 'recom_familiares', 'recom_ejercicio', 'recom_iglesia', 'haceds_hogar', 'haceds_vecinos', 'haceds_jovenes', 'haceds_mayores', 'haceds_otropais', 'haceds_emplmercados', 'medios_noti_redessociales', 'medios_noti_chat', 'medios_noti_periodicos', 'medios_noti_tv', 'medios_noti_radio', 'medios_covid_redessociales', 'medios_covid_chat', 'medios_covid_periodicos', 'medios_covid_tv', 'medios_covid_radio', 'nps_medio_redessoc', 'nps_medio_sms', 'nps_medio_periodico', 'nps_medio_tv', 'nps_medio_radio', 'nps_medio_website', 'nps_medio_otro']
tipo_columnas = ['CategoricalDtype', 'dtype[object_]', 'dtype[float32]', 'dtype[float32]', 'dtype[float64]', 'dtype[object_]', 'CategoricalDtype', 'dtype[int16]', 'CategoricalDtype', 'CategoricalDtype', 'dtype[int8]', 'dtype[int8]', 'CategoricalDtype', 'CategoricalDtype', 'dtype[float64]', 'dtype[float64]', 'CategoricalDtype', 'dtype[float64]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'dtype[float32]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'dtype[float64]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'dtype[float64]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'dtype[int8]', 'dtype[int8]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype']

# Caso 1: no existe la función.
try:
    source = inspect.getsource(indexar_columnas_numericamente)
    assert type(indexar_columnas_numericamente) == type(lambda:None)
except:
    raise NotImplementedError("No existe una función llamada indexar_columnas_numericamente.",)
    
# Caso 2: la función es interrumpida por errores durante su ejecución.
try:
    resultado = indexar_columnas_numericamente()
except:
    raise RuntimeError("Tu función produce un error al ejecutarse.")

# Caso 3: no retorna un DataFrame
assert type(resultado) == pd.DataFrame, f"Tu función debe retornar un valor de tipo {pd.DataFrame.__name__}."

# Caso 4: no ejecuta la función importar_e_indexar
assert source.find("importar_e_indexar()") >= 0, "Tu código no hace referencia a la función importar_e_indexar. Asegurate de ejecutarla en tu código."

# Caso 5: retorna el dataframe con la columna 'id'
assert "id" not in resultado.columns, "Tu función retorna un DataFrame que incluye la columna 'id'. Esta columna debes convertirla en el index del DataFrame."

# Caso 6: retorna un dataframe con cantidad de columnas errada
assert resultado.shape[1] == 107, "Tu función retorna un DataFrame con cantidad de columnas errada."

# Caso 7: devuelve un dataframe cuyas columnas tienen nombre distinto de lo esperado
assert [*resultado.columns] == nombre_columnas, "Tu función retorna un DataFrame con columnas cuyos nombres no coinciden con los de los datos del archivo. Asegurate de estar filtrando las columnas indicadas."

# Caso 8: devuelve un dataframe cuyas columnas tienen tipo distinto de lo esperado
assert [i.__class__.__name__ for i in resultado.dtypes] == tipo_columnas, "Tu función retorna un DataFrame con columnas cuyos tipos no coinciden con los de los datos del archivo. Asegurate de estar filtrando las columnas indicadas."

# Caso 9: devuelve un dataframe con cantidad de filas errada
assert resultado.shape[0] == 230509, "Tu función retorna un DataFrame con cantidad de filas errada."

# Caso 10: devuelve un dataframe con indice del mismo tipo que por defecto ('int64')
assert resultado.index.dtype.name != 'int64', "Tu función retorna un DataFrame con índice del mismo tipo que el que se crea por defecto en pandas. Asegurate de estar asignando la columna 'id' como índice."

# Caso 11: devuelve un dataframe con índice de tipo errado
assert resultado.index.dtype.name == 'float64', "Tu función retorna un DataFrame con índice de tipo distinto del esperado. Asegurate de estar usando la columna 'id' sin editarla."

# Caso 12: primer o última posicion del índice distintas de lo esperado
assert resultado.index[0] == 1000060.0, "Tu función retorna un DataFrame que es errado por lo menos en la primera posición de su índice"
assert resultado.index[len(resultado.index) - 1] == 17003990.0, "Tu función retorna un DataFrame que es errado por lo en menos la última posición de su índice"


# Mensaje de felicitaciones
print("Felicidades, realizaste este ejercicio correctamente.")
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-12-337198a35f7b> in <module>()
     36 
     37 # Caso 8: devuelve un dataframe cuyas columnas tienen tipo distinto de lo esperado
---> 38 assert [i.__class__.__name__ for i in resultado.dtypes] == tipo_columnas, "Tu función retorna un DataFrame con columnas cuyos tipos no coinciden con los de los datos del archivo. Asegurate de estar filtrando las columnas indicadas."
     39 
     40 # Caso 9: devuelve un dataframe con cantidad de filas errada

AssertionError: Tu función retorna un DataFrame con columnas cuyos tipos no coinciden con los de los datos del archivo. Asegurate de estar filtrando las columnas indicadas.
```

### Ejercicio 5

A raiz de la pandemia, los gobiernos de latinoamérica desean monitorear a aquellas personas que salen de sus casas a trabajar y conviven con personas mayores de 60 años.

Implementa una función llamada `trabajo_afuera_y_mayores` que no reciba parámetros, tome el `DataFrame` resultante del [Ejercicio 1](#Ejercicio-1) y lo retorne únicamente con los registros que en la columna `"trabajocasa"` tengan el valor `"Fuera de su casa"` y que en la columna `"haymayores"` tengan el valor `"si"`.

La función debe retornar un `DataFrame`.

<!-- grade_id: ej5-sol, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def trabajo_afuera_y_mayores():
    df5 = importar_e_indexar()
    #Filtro de columnas del DF 
    return df5[(df5["trabajocasa"] == "Fuera de su casa") & (df5["haymayores"] == "si")]
```

<!-- grade_id: ej5-test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Base variables
import inspect

nombre_columnas = ['country_code', 'country_name', 'weight', 'weight_pop', 'pop', 'recordeddate', 'genero', 'edad', 'educ', 'estcivil', 'nrhogar', 'nrhogar_antes', 'hayninhos', 'ninos_5', 'nrninhos', 'matriculados', 'esc_publica', 'grado', 'haymayores', 'diasmercado', 'lugarmercado', 'comprotodo', 'nocompra_pan', 'nocompra_vegetal', 'nocompra_carne', 'nocompra_envasados', 'nocompra_bebidas', 'nocompra_limpieza', 'nocompra_medicinas', 'nocompra_otro', 'nrnegocios', 'esperarfila', 'subaprecios', 'hambre', 'menos_saludable', 'vivienda', 'probvivienda', 'probvivienda_prox', 'ayuda_alquiler', 'sufialimento', 'sufirecursos', 'gasto_emergencia', 'monto_gasto', 'pre_cash', 'new_cash', 'coima', 'tiene_banco', 'tiene_tarjdebito', 'tiene_segsaludpub', 'tiene_segsaludpriv', 'recibeprestamos', 'darprestamos', 'remesas_enero', 'remesas_origen', 'remesas_persona', 'recibio_remesas', 'lm_perdioempleo', 'lm_cerronegocio_gob', 'lm_cerronegocio_dem', 'lm_asistmedica', 'lm_perdioingrarriendo', 'lm_dejoasistencia', 'lm_freq', 'recibio_ayudaemp', 'compromiso_futurotrabajo', 'empleo_formal', 'actividad_empleado', 'actividad_negocio', 'actividad_estudio', 'trabajocasa', 'diastrabajo', 'horastrabajo', 'sector', 'redujo_ingreso', 'ingreso_jan2020', 'ingreso_apr2020', 'retorno_ingreso', 'sintoma_ninguno', 'sintoma_garganta', 'sintoma_gripe', 'sintoma_sarpullido', 'sintoma_fiebre', 'sintoma_tos', 'conoce_sintomas', 'infectados_sinsintomas', 'sinsintomas_contagiar', 'visito_trabajo', 'visito_hospital', 'visito_mercado', 'visito_banco', 'visito_amigos', 'visito_familiares', 'visito_ejercitar', 'visito_iglesia', 'comotransmite', 'sabe_comotransmite', 'distsoc', 'recom_trabajo', 'recom_hospital', 'recom_mercado', 'recom_banco', 'recom_amigos', 'recom_familiares', 'recom_ejercicio', 'recom_iglesia', 'haceds_hogar', 'haceds_vecinos', 'haceds_jovenes', 'haceds_mayores', 'haceds_otropais', 'haceds_emplmercados', 'medios_noti_redessociales', 'medios_noti_chat', 'medios_noti_periodicos', 'medios_noti_tv', 'medios_noti_radio', 'medios_covid_redessociales', 'medios_covid_chat', 'medios_covid_periodicos', 'medios_covid_tv', 'medios_covid_radio', 'nps_medio_redessoc', 'nps_medio_sms', 'nps_medio_periodico', 'nps_medio_tv', 'nps_medio_radio', 'nps_medio_website', 'nps_medio_otro']
tipo_columnas = ['CategoricalDtype', 'dtype[object_]', 'dtype[float32]', 'dtype[float32]', 'dtype[float64]', 'dtype[object_]', 'CategoricalDtype', 'dtype[int16]', 'CategoricalDtype', 'CategoricalDtype', 'dtype[int8]', 'dtype[int8]', 'CategoricalDtype', 'CategoricalDtype', 'dtype[float64]', 'dtype[float64]', 'CategoricalDtype', 'dtype[float64]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'dtype[float32]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'dtype[float64]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'dtype[float64]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'dtype[int8]', 'dtype[int8]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'dtype[float32]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'dtype[object_]', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype', 'CategoricalDtype']

# Caso 1: no existe la función.
try:
    source = inspect.getsource(trabajo_afuera_y_mayores)
    assert type(trabajo_afuera_y_mayores) == type(lambda:None)
except:
    raise NotImplementedError("No existe una función llamada trabajo_afuera_y_mayores.",)
    
# Caso 2: la función es interrumpida por errores durante su ejecución.
try:
    resultado = trabajo_afuera_y_mayores()
except:
    raise RuntimeError("Tu función produce un error al ejecutarse.")

# Caso 3: no retorna un DataFrame
assert type(resultado) == pd.DataFrame, f"Tu función debe retornar un valor de tipo {pd.DataFrame.__name__}."

# Caso 4: no ejecuta la función importar_e_indexar
assert source.find("importar_e_indexar()") >= 0, "Tu código no hace referencia a la función importar_e_indexar. Asegurate de ejecutarla en tu código."

# Caso 5: retorna el dataframe con la columna 'id'
assert "id" not in resultado.columns, "Tu función retorna un DataFrame que incluye la columna 'id'. Esta columna debes haberla convertido en el index sencillo del DataFrame."

# Caso 6: retorna un dataframe con cantidad de columnas errada
assert resultado.shape[1] == 128, "Tu función retorna un DataFrame con cantidad de columnas errada."

# Caso 7: devuelve un dataframe cuyas columnas tienen nombre distinto de lo esperado
assert [*resultado.columns] == nombre_columnas, "Tu función retorna un DataFrame con columnas cuyos nombres no coinciden con los de los datos del archivo."

# Caso 8: devuelve un dataframe cuyas columnas tienen tipo distinto de lo esperado
assert [i.__class__.__name__ for i in resultado.dtypes] == tipo_columnas, "Tu función retorna un DataFrame con columnas cuyos tipos no coinciden con los de los datos del archivo."

# Caso 9: devuelve un dataframe con cantidad de filas errada
assert resultado.shape[0] == 14311, "Tu función retorna un DataFrame con cantidad de filas errada."

# Caso 10: devuelve un dataframe con indice del mismo tipo que por defecto ('int64')
assert resultado.index.dtype.name != 'int64', "Tu función retorna un DataFrame con índice del mismo tipo que el que se crea por defecto en pandas. Asegurate de estar asignando la columna 'id' como índice."

# Caso 11: devuelve un dataframe con índice de tipo errado
assert resultado.index.dtype.name == 'float64', "Tu función retorna un DataFrame con índice de tipo distinto del esperado. Asegurate de estar usando la columna 'id' sin editarla."

# Caso 12: primera o última posicion del índice distintas de lo esperado
assert resultado.index[0] == 1000019.0, "Tu función retorna un DataFrame que es errado por lo menos en la primera posición de su índice"
assert resultado.index[len(resultado.index) - 1] == 17020778.0, "Tu función retorna un DataFrame que es errado por lo menos en la última posición de su índice"


# Mensaje de felicitaciones
print("Felicidades, realizaste este ejercicio correctamente.")
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-17-6845f1b2f3ae> in <module>()
     36 
     37 # Caso 8: devuelve un dataframe cuyas columnas tienen tipo distinto de lo esperado
---> 38 assert [i.__class__.__name__ for i in resultado.dtypes] == tipo_columnas, "Tu función retorna un DataFrame con columnas cuyos tipos no coinciden con los de los datos del archivo."
     39 
     40 # Caso 9: devuelve un dataframe con cantidad de filas errada

AssertionError: Tu función retorna un DataFrame con columnas cuyos tipos no coinciden con los de los datos del archivo.
```

### Ejercicio 6

Implementa una función llamada `hogares_decrecen` que no reciba parámetros, tome el `DataFrame` resultante del [Ejercicio 1](#Ejercicio-1) y retorne un `DataFrame` únicamente con las columnas `"nrhogar"`, `"nrhogar_antes"` y una columna adicional llamada `"hogar_decrecio"` que para cada registro tome un valor de:

* `1`, si el valor numérico de la columna `"nrhogar"` es menor que el valor de la columna `"nrhogar_antes"`;

* `0`, de lo contrario.

La función debe retornar un `DataFrame`.

<!-- grade_id: ej6-sol, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def hogares_decrecen():
    df6 = importar_e_indexar()
    df6 = df6[["nrhogar", "nrhogar_antes"]]
    df6["hogar_decrecio"] = (df6["nrhogar"] < df6["nrhogar_antes"]) * 1
    return df6
```

<!-- grade_id: ej6-test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Base variables
import inspect

nombre_columnas = ['nrhogar','nrhogar_antes','hogar_decrecio']
tipo_columnas = ['int', 'int', 'int']

# Caso 1: no existe la función.
try:
    source = inspect.getsource(hogares_decrecen)
    assert type(hogares_decrecen) == type(lambda:None)
except:
    raise NotImplementedError("No existe una función llamada hogares_decrecen.",)
    
# Caso 2: la función es interrumpida por errores durante su ejecución.
try:
    resultado = hogares_decrecen()
except:
    raise RuntimeError("Tu función produce un error al ejecutarse.")

# Caso 3: no retorna un DataFrame
assert type(resultado) == pd.DataFrame, f"Tu función debe retornar un valor de tipo {pd.DataFrame.__name__}."

# Caso 4: no ejecuta la función importar_e_indexar
assert source.find("importar_e_indexar()") >= 0, "Tu código no hace referencia a la función importar_e_indexar. Asegurate de ejecutarla en tu código."

# Caso 5: retorna el dataframe con la columna 'id'
assert "id" not in resultado.columns, "Tu función retorna un DataFrame que incluye la columna 'id'. Esta columna debes convertirla en el index del DataFrame."

# Caso 6: retorna un dataframe con cantidad de columnas errada
assert resultado.shape[1] == 3, "Tu función retorna un DataFrame con cantidad de columnas errada."

# Caso 7: devuelve un dataframe cuyas columnas tienen nombre distinto de lo esperado
assert [*resultado.columns] == nombre_columnas, "Tu función retorna un DataFrame con columnas cuyos nombres no son nrhogar, nrhogar_antes y hogar_decrecio (en ese orden)."

# Caso 8: devuelve un dataframe cuyas columnas tienen tipo distinto de lo esperado
assert [i.name[0:3] for i in resultado.dtypes] == tipo_columnas, "Tu función retorna un DataFrame con columnas cuyos tipos no coinciden con los esperados. Asegurate de registrar los 0 y 1 como números enteros."

# Caso 9: devuelve un dataframe con cantidad de filas errada
assert resultado.shape[0] == 230509, "Tu función retorna un DataFrame con cantidad de filas errada."

# Caso 10: devuelve un dataframe con indice del mismo tipo que por defecto ('int64')
assert resultado.index.dtype.name != 'int64', "Tu función retorna un DataFrame con índice del mismo tipo que el que se crea por defecto en pandas. Asegurate de estar asignando la columna 'id' como índice."

# Caso 11: devuelve un dataframe con índice de tipo errado
assert resultado.index.dtype.name == 'float64', "Tu función retorna un DataFrame con índice de tipo distinto del esperado. Asegurate de estar usando la columna 'id' sin editarla."

# Caso 12: primer o última posicion del índice distintas de lo esperado
assert resultado.index[0] == 1000060.0, "Tu función retorna un DataFrame que es errado por lo menos en la primera posición de su índice"
assert resultado.index[len(resultado.index) - 1] == 17003990.0, "Tu función retorna un DataFrame que es errado por lo menos en la última posición de su índice"

# Caso 13: asigna más de dos valores únicos a la columna hogar_decrecio
assert len(resultado["hogar_decrecio"].unique()) == 2, "Tu función retorna un DataFrame que que en la columna hogar_decrecio incluye más de dos valores únicos. Recuerda que esta columna solo debe incluir los números enteros 0 y 1."

# Caso 14: no asigna 0 y 1 como únicos valores a la columna hogar_decrecio
assert 0 in resultado["hogar_decrecio"].unique() and 1 in resultado["hogar_decrecio"].unique(), "Tu función retorna un DataFrame que que en la columna hogar_decrecio incluye valores distintos de los números enteros 0 y 1."

# Caso 15: asigna mal el valor binario.
assert sum(resultado["hogar_decrecio"]) == 19908, "Tu función retorna un DataFrame que que en la columna hogar_decrecio asigna de manera errada los valores 0 y 1."


# Mensaje de felicitaciones
print("Felicidades, realizaste este ejercicio correctamente.")
```

Output:
```
Felicidades, realizaste este ejercicio correctamente.
```

## Referencias

BID (2020). Encuesta Coronavirus BID/Cornell. Recuperado el 11 de febrero de 2020 de: https://data.iadb.org/DataCatalog/Dataset#DataCatalogID=11319/28452

## Créditos

**Autor(es)**: Jorge Esteban Camargo Forero, Alejandro Mantilla Redondo, Diego Alejandro Cely Gómez

**Fecha última actualización:** 10/09/2021